# Q2: Library management system

Imagine you are building a simple library management system. Implement three classes: `Author`, `Book`, and `Library`.

The **Author** class should have the following attributes:
- `name (str)`: the name of the author
- `birth_year (int)`: the birth year of the author

It should also have a method:
- `display_info()`: prints out the author's name and birth year.

The **Book** class should have the following attributes:
- `title (str)`: the title of the book
- `isbn (str)`: the International Standard Book Number of the book
- `author (Author)`: the author of the book (an instance of the Author class)
- `available_copies (int)`: the number of available copies of the book

It should also have the following methods:
- `display_info()`: prints out the title, ISBN, author's name, and the number of available copies of the book.
- `borrow_copy()`: Decreases the number of available copies by 1 when a book is borrowed. If no copies are available, print a message indicating that the book is currently unavailable.
- `return_copy()`: Increases the number of available copies by 1 when a book is returned.

The **Library** class should have the following attributes:
- `books(dict)`: a list containing instances of the Book class

It should also have the following methods:
- `add_book(book: Book)`: adds a book to the library's collection
- `display_books()`: displays information about all the books in the library
- `search_books_by_author(author_name: str)`: Takes an author's name as input and prints the titles of all books in the library written by that author.
- `search_books_by_title(title: str)`: Takes a title as input and prints information about the book with that title in the library.
- `total_available_copies()`: Calculates and returns the total number of available copies for all books in the library.

Implement the classes and methods described above and create an example demonstrating the use of class composition with these classes.

<!-- grade_id: cell-832ef3d3172fac05, points: 0, grade: false, solution: true -->
```python
class Author:
    def __init__(self, name, birth_year):
        # YOUR CODE HERE
        self.name = name
        self.birth_year = birth_year

    def display_info(self):
        # YOUR CODE HERE
        return f"Name : {self.name} , Birth_year : {self.birth_year}"
class Book:
    def __init__(self, title, isbn, author, available_copies):
        # YOUR CODE HERE
        self.title = title
        self.isbn = isbn
        self.author = author
        self.available_copies = available_copies

    def display_info(self):
        # YOUR CODE HERE
        return f"Title: {self.title}, ISBN: {self.isbn}, Author: {self.author.name}, Available Copies: {self.available_copies}"
    def borrow_copy(self):
        # YOUR CODE HERE
        if self.available_copies <= 0:
            return f'ฺBook currently unavailable'
        else:
            self.available_copies -= 1
            return f'Book borrowed successfully'
        
    def return_copy(self):
        # YOUR CODE HERE
        self.available_copies += 1
        return f'Book returned successfully'
        
class Library:
    def __init__(self):
        self.books = {}
        
    def add_book(self, book):
        # YOUR CODE HERE
        self.books[book.isbn] = book

    def display_books(self):
        book_details = ""
        for book in self.books.values():
            book_details += book.display_info() + "\n"  
        return book_details.strip()


    def search_books_by_author(self, author_name):
        # YOUR CODE HERE
        author_book = ''
        for book in self.books.values():
            if book.author.name == author_name:
                author_book += book.display_info() + '\n'
            if not author_book:
                return f'No books found'
        return author_book.strip()

    def search_books_by_title(self, title):
        # YOUR CODE HERE
        title_book = ''
        for book in self.books.values():
            if book.title == title:
                title_book += book.display_info() + '\n'
            if not title_book:
                return f'No book found'
        return title_book.strip()
    def total_available_copies(self):
        # YOUR CODE HERE
        total_copies = sum(book.available_copies for book in self.books.values())
        return total_copies
        
        
```

In the context of Python's unittest framework, each test case is executed in a new instance of the test case class. When a test case is run, the `setUp` method is called before each test method, and the `tearDown` method is called after each test method. 

**Author Class Test Cases:**
- Test Case 1.1: Create an Author instance with a valid name and birth year.
- Test Case 1.2: Verify that display_info() method returns a string.

<!-- grade_id: cell-f2e38ae36a6170a5, points: 1, grade: true, solution: false -->
```python
import unittest

class TestAuthor(unittest.TestCase):
    
    def setUp(self): 
        self.author = Author("John Doe", 1980)

    def test_author_creation(self):
        # Test Case 1.1
        self.assertIsInstance(self.author, Author)
        print("ok")

    def test_author_display_info(self):
        # Test Case 1.2
        author_info = self.author.display_info()
        self.assertIsInstance(author_info, str)
        self.assertIn("John Doe", author_info)
        self.assertIn("1980", author_info)
        print("ok")

# Load and run TestAuthor
result = unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestAuthor))
if len(result.failures) == 0:
    pass
```

Output:
```
..
----------------------------------------------------------------------
Ran 2 tests in 0.003s

OK
ok
ok
```

**Book Class Test Cases:**
- Test Case 2.1: Create a Book instance with a valid title, ISBN, Author instance, and available copies.
- Test Case 2.2: Verify that display_info() method returns a string.
- Test Case 2.3: Borrow a copy of a book with available copies.
- Test Case 2.4: Return a copy of a book.
- Test Case 2.5: Attempt to borrow a copy of a book with no available copies.

<!-- grade_id: cell-165c2239435398d8, points: 1, grade: true, solution: false -->
```python
import unittest

class TestBook(unittest.TestCase):    
  
    def setUp(self): 
        self.author = Author("John Doe", 1980)
        self.book1 = Book("Introduction to Python", "978-1-234567-89-0", self.author, 10)
        self.book2 = Book("Data Structures in Python", "978-1-234567-89-1", self.author, 5)
        self.book3 = Book("Advanced Python Programming", "978-1-234567-89-2", self.author, 1)
        
    def test_book_creation(self):
        # Test Case 2.1
        self.assertIsInstance(self.book1, Book)
        print("ok")

    def test_book_display_info(self):
        # Test Case 2.2
        book_info = self.book1.display_info()
        self.assertIsInstance(book_info, str)
        self.assertIn("Introduction to Python", book_info)
        self.assertIn("978-1-234567-89-0", book_info)
        self.assertIn("John Doe", book_info)
        self.assertIn("10", book_info)
        print("ok")

    def test_borrow_copy(self):
        # Test Case 2.3
        borrow_result = self.book1.borrow_copy()
        self.assertIsInstance(borrow_result, str)
        self.assertIn("borrowed successfully", borrow_result)
        print("ok")

    def test_return_copy(self):
        # Test Case 2.4
        return_result = self.book2.return_copy()
        self.assertIsInstance(return_result, str)
        self.assertIn("returned successfully", return_result)
        print("ok")

    def test_borrow_copy_no_available_copies(self):
        # Test Case 2.5
        self.book3.borrow_copy()  # Borrow the only available copy
        borrow_result = self.book3.borrow_copy()
        self.assertIsInstance(borrow_result, str)
        self.assertIn("currently unavailable", borrow_result)
        print("ok")
        
# Load and run TestBook
result = unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestBook))
if len(result.failures) == 0:
    pass
```

Output:
```
.....
----------------------------------------------------------------------
Ran 5 tests in 0.119s

OK
ok
ok
ok
ok
ok
```

**Library Class Test Cases:**
- Test Case 3.1: Create a Library instance.
- Test Case 3.2: Add a Book instance to the Library.
- Test Case 3.3: Verify that display_books() method returns a string.
- Test Case 3.4: Search for books by a valid author's name.
- Test Case 3.5: Search for a book by a valid title.
- Test Case 3.6: Attempt to search for books by an invalid author's name.
- Test Case 3.7: Attempt to search for a non-existing book by title.
- Test Case 3.8: Verify that total_available_copies() method returns an integer.

<!-- grade_id: cell-95555308a475811e, points: 1, grade: true, solution: false -->
```python
import unittest

class TestLibrary(unittest.TestCase):
  
    def setUp(self): 
        self.author = Author("John Doe", 1980)
        self.book1 = Book("Introduction to Python", "978-1-234567-89-0", self.author, 10)
        self.book2 = Book("Data Structures in Python", "978-1-234567-89-1", self.author, 5)
        self.book3 = Book("Advanced Python Programming", "978-1-234567-89-2", self.author, 1)
        self.library = Library()
        self.library.add_book(self.book1)
        self.library.add_book(self.book2)
        self.library.add_book(self.book3)
        
    def test_library_creation(self):
        # Test Case 3.1
        self.assertIsInstance(self.library, Library)
        print("ok")

    def test_add_book(self):
        # Test Case 3.2
        self.assertEqual(len(self.library.books), 3)
        print("ok")

    def test_display_books(self):
        # Test Case 3.3
        library_info = self.library.display_books()
        self.assertIsInstance(library_info, str)
        self.assertIn("Introduction to Python", library_info)
        self.assertIn("Data Structures in Python", library_info)
        self.assertIn("Advanced Python Programming", library_info)
        print("ok")

    def test_search_books_by_author(self):
        # Test Case 3.4
        author_books = self.library.search_books_by_author("John Doe")
        self.assertIsInstance(author_books, str)
        self.assertIn("Introduction to Python", author_books)
        print("ok")

    def test_search_books_by_title(self):
        # Test Case 3.5
        book_info = self.library.search_books_by_title("Introduction to Python")
        self.assertIsInstance(book_info, str)
        self.assertIn("Introduction to Python", book_info)
        print("ok")

    def test_search_books_by_invalid_author(self):
        # Test Case 3.6
        not_found_msg = self.library.search_books_by_author("Unknown Author")
        self.assertIsInstance(not_found_msg, str)
        self.assertIn("No books found", not_found_msg)
        print("ok")

    def test_search_books_by_non_existing_title(self):
        # Test Case 3.7
        not_found_msg = self.library.search_books_by_title("Non-Existing Title")
        self.assertIsInstance(not_found_msg, str)
        self.assertIn("No book found", not_found_msg)
        print("ok")

    def test_total_available_copies(self):
        # Test Case 3.8
        isbn = "978-1-234567-89-0"
        self.library.books[isbn].borrow_copy()
        self.library.books[isbn].borrow_copy()
        self.library.books[isbn].return_copy()
        total_copies = self.library.total_available_copies()
        self.assertIsInstance(total_copies, int)
        self.assertEqual(total_copies, 15)
        print("ok")
        
# Load and run TestLibrary
result = unittest.TextTestRunner().run(unittest.TestLoader().loadTestsFromTestCase(TestLibrary))
if len(result.failures) == 0:
    pass
```

Output:
```
........
----------------------------------------------------------------------
Ran 8 tests in 0.016s

OK
ok
ok
ok
ok
ok
ok
ok
ok
```

```python
# Example Usage
author1 = Author("John Doe", 1980)
book1 = Book("Introduction to Python", "978-1-234567-89-0", author1, 10)
book2 = Book("Data Structures in Python", "978-1-234567-89-1", author1, 5)
book3 = Book("Advanced Python Programming", "978-1-234567-89-2", author1, 1)


library = Library()
library.add_book(book1)
library.add_book(book2)
library.add_book(book3)

# Display information about the books in the library
print(library.display_books())

# Borrow a copy of a book
print(library.books[book1.isbn].borrow_copy())

# Return a copy of a book
print(library.books[book2.isbn].return_copy())

# Display updated information about the books in the library
print(library.display_books())

# Search for books by a specific author
print(library.search_books_by_author("John Doe"))

# Search for a book by title
print(library.search_books_by_title("Introduction to Python"))

# Get the total number of available copies in the library
print(library.total_available_copies())

```

Output:
```
Title: Introduction to Python, ISBN: 978-1-234567-89-0, Author: John Doe, Available Copies: 10
Title: Data Structures in Python, ISBN: 978-1-234567-89-1, Author: John Doe, Available Copies: 5
Title: Advanced Python Programming, ISBN: 978-1-234567-89-2, Author: John Doe, Available Copies: 1
Book borrowed successfully
Book returned successfully
Title: Introduction to Python, ISBN: 978-1-234567-89-0, Author: John Doe, Available Copies: 9
Title: Data Structures in Python, ISBN: 978-1-234567-89-1, Author: John Doe, Available Copies: 6
Title: Advanced Python Programming, ISBN: 978-1-234567-89-2, Author: John Doe, Available Copies: 1
Title: Introduction to Python, ISBN: 978-1-234567-89-0, Author: John Doe, Available Copies: 9
Title: Data Structures in Python, ISBN: 978-1-234567-89-1, Author: John Doe, Available Copies: 6
Title: Advanced Python Programming, ISBN: 978-1-234567-89-2, Author: John Doe, Available Copies: 1
Title: Introduction to Python, ISBN: 978-1-234567-89-0, Author: John Doe, A
```